cnf(c_0_1,negated_conjecture,(a|b), file('../../../iprover/test.p', clause_1)).
cnf(c_0_2,negated_conjecture,(~a|~b), file('../../../iprover/test.p', clause_2)).
cnf(c_0_3,negated_conjecture,(a|c|~d), file('../../../iprover/test.p', clause_3)).
cnf(c_0_4,negated_conjecture,(a|d|~c), file('../../../iprover/test.p', clause_4)).
cnf(c_0_5,negated_conjecture,(c|d|~a), file('../../../iprover/test.p', clause_5)).
cnf(c_0_6,negated_conjecture,(~a|~c|~d), file('../../../iprover/test.p', clause_6)).
cnf(c_0_7,negated_conjecture,(b|c|~e), file('../../../iprover/test.p', clause_7)).
cnf(c_0_8,negated_conjecture,(b|e|~c), file('../../../iprover/test.p', clause_8)).
cnf(c_0_9,negated_conjecture,(c|e|~b), file('../../../iprover/test.p', clause_9)).
cnf(c_0_10,negated_conjecture,(~b|~c|~e), file('../../../iprover/test.p', clause_10)).
cnf(c_0_11,negated_conjecture,(d|~e), file('../../../iprover/test.p', clause_11)).
cnf(c_0_12,negated_conjecture,(e|~d), file('../../../iprover/test.p', clause_12)).

# Auto-Ordering is analysing problem.
# Problem is type GHNFGFFSF00SS
# Auto-mode selected ordering type KBO6
# Auto-mode selected ordering precedence scheme <invfreq>
# Auto-mode selected weight ordering scheme <invfreqrank>
#
# Auto-Heuristic is analysing problem.
# Problem is type GHNFGFFSF00SS
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# No equality, disabling AC handling.
#
# Initializing proof state
# Presaturation interreduction done
cnf(c_0_13,negated_conjecture,(d|~d),inference(spm,[status(thm)],[c_0_11,c_0_12,theory(equality)])).
cnf(c_0_14,negated_conjecture,(b|~b),inference(spm,[status(thm)],[c_0_2,c_0_1,theory(equality)])).
cnf(c_0_15,negated_conjecture,(b|~d|~c),inference(spm,[status(thm)],[c_0_6,c_0_1,theory(equality)])).
cnf(c_0_16,negated_conjecture,(c|~b|~d),inference(spm,[status(thm)],[c_0_2,c_0_3,theory(equality)])).
cnf(c_0_17,negated_conjecture,(c|~d|~c),inference(spm,[status(thm)],[c_0_6,c_0_3,theory(equality)])).
cnf(c_0_18,negated_conjecture,(~c|~b|~d),inference(spm,[status(thm)],[c_0_10,c_0_12,theory(equality)])).
cnf(c_0_19,negated_conjecture,(~d|~c),inference(csr,[status(thm)],[c_0_15,c_0_18,theory(equality,[symmetry])])).
cnf(c_0_20,negated_conjecture,(~d|~b),inference(csr,[status(thm)],[c_0_16,c_0_19,theory(equality,[symmetry])])).
cnf(c_0_21,negated_conjecture,(a|~c),inference(csr,[status(thm)],[c_0_4,c_0_19,theory(equality,[symmetry])])).
cnf(c_0_22,negated_conjecture,(~b|~c),inference(spm,[status(thm)],[c_0_2,c_0_21,theory(equality)])).
cnf(c_0_23,negated_conjecture,(c|b|~d),inference(spm,[status(thm)],[c_0_7,c_0_12,theory(equality)])).
cnf(c_0_24,negated_conjecture,(d|c|b),inference(spm,[status(thm)],[c_0_5,c_0_1,theory(equality)])).
cnf(c_0_25,negated_conjecture,(d|c|~d),inference(spm,[status(thm)],[c_0_5,c_0_3,theory(equality)])).
cnf(c_0_26,negated_conjecture,(d|c|~c),inference(spm,[status(thm)],[c_0_5,c_0_21,theory(equality)])).
cnf(c_0_27,negated_conjecture,(c|~d),inference(csr,[status(thm)],[c_0_23,c_0_20,theory(equality,[symmetry])])).
cnf(c_0_28,negated_conjecture,(~d),inference(csr,[status(thm)],[c_0_27,c_0_19,theory(equality,[symmetry])])).
cnf(c_0_29,negated_conjecture,(e|~c),inference(csr,[status(thm)],[c_0_8,c_0_22,theory(equality,[symmetry])])).
cnf(c_0_30,negated_conjecture,(d|~c),inference(spm,[status(thm)],[c_0_11,c_0_29,theory(equality)])).
cnf(c_0_31,negated_conjecture,(c|b|~c),inference(spm,[status(thm)],[c_0_7,c_0_29,theory(equality)])).
cnf(c_0_32,negated_conjecture,(~c),inference(sr,[status(thm)],[c_0_30,c_0_28,theory(equality)])).
cnf(c_0_33,negated_conjecture,(e|~b),inference(csr,[status(thm)],[c_0_9,c_0_22,theory(equality,[symmetry])])).
cnf(c_0_34,negated_conjecture,(d|~b),inference(spm,[status(thm)],[c_0_11,c_0_33,theory(equality)])).
cnf(c_0_35,negated_conjecture,(c|b|~b),inference(spm,[status(thm)],[c_0_7,c_0_33,theory(equality)])).
cnf(c_0_36,negated_conjecture,(~b),inference(sr,[status(thm)],[c_0_34,c_0_28,theory(equality)])).
cnf(c_0_37,negated_conjecture,(a),inference(sr,[status(thm)],[c_0_1,c_0_36,theory(equality)])).
cnf(c_0_38,negated_conjecture,(d|c|$false),inference(rw, [status(thm)],[c_0_5,c_0_37,theory(equality)])).
cnf(c_0_39,negated_conjecture,(d|c),inference(cn,[status(thm)],[c_0_38, theory(equality,[symmetry])])).
cnf(c_0_40,negated_conjecture,(c),inference(sr,[status(thm)],[c_0_39,c_0_28,theory(equality)])).
cnf(c_0_41,negated_conjecture,($false),inference(sr,[status(thm)],[c_0_40,c_0_32,theory(equality)])).
cnf(c_0_42,negated_conjecture,($false), c_0_41,['proof']).

# Proof found!
# SZS status Unsatisfiable
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 12
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 12
# Processed clauses                    : 32
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 32
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 5
# Backward-rewritten                   : 2
# Generated clauses                    : 16
# ...of the previous two non-trivial   : 9
# Contextual simplify-reflections      : 7
# Paramodulations                      : 15
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 12
#    Positive orientable unit clauses  : 1
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 8
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Clause-clause subsumption calls (NU) : 18
# Rec. Clause-clause subsumption calls : 10
# Non-unit clause-clause subsumptions  : 10
# Unit Clause-clause subsumption calls : 15
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
